```python

```

<!-- grade_id: cell-9c504fcbd98c74bd, points: 0, grade: false, solution: true -->
```python
x = ...
# your code here

```

<!-- grade_id: cell-7f3fe7aed80e1fe1, points: 1, grade: true, solution: false -->
```python
# проверка, просто запустите ячейку

assert x+1 == 2
```

```python

```